Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name below:

```python
NAME = "Ido Rabin" 

```

---

# Assignment 2: Applications of modular arithmetics

```python
import math
```

## Question 1: Cryptography

You will implement and analyze two simple cryptographic algorithms — Caesar's and Scytale ciphers.

1. Implement the encoder for Ceasar's cipher: https://en.wikipedia.org/wiki/Caesar_cipher (encode only the letters, leave other characters as-is).

Q1.1: Implement function `ceasar_cipher`, which will transform regular text into encoded text
<pre>
------------
Input parameters:
- plaintext - string, the text that should be encrypted
- N - int, size of shift
------------
return value:
string, the text after encoding
</pre>

<!-- grade_id: cell-fcc26ea16c7fcc6e, points: 0, grade: false, solution: true -->
```python
def ceasar_cipher(plaintext, shift_amount):
    """
    Encrypts the given plaintext using the Caesar cipher algorithm with the specified shift amount.

    Args:
    plaintext (str): The text to be encrypted.
    shift_amount (int): The number of positions each character in the plaintext will be shifted.

    Returns:
    str: The encrypted text (ciphertext).
    """
    ciphertext = ""
    alphabet = "abcdefghijklmnopqrstuvwxyz"

    for char in plaintext:
        if char.lower() in alphabet:
            original_index = alphabet.find(char.lower())
            new_index = (original_index + shift_amount) % len(alphabet)
            ciphertext += alphabet[new_index]
        else:
            ciphertext += char

    return ciphertext

```

<!-- grade_id: cell-910c1caba85f9efe, points: 10, grade: true, solution: false -->
```python
# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q1.1 --- Test your implementation:
# ---------------------------
print ("Testing the implementation of the 'ceasarCipher' function..\n ")
assert type(ceasar_cipher("meep meep",1))==str,'output type is incorrect, should be string'
print ("good job!\nTests passed. There may be hidden tests too")

```

Output:
```
Testing the implementation of the 'ceasarCipher' function..
 
good job!
Tests passed. There may be hidden tests too
```

2. Implement a decoder function that can reverse the cypher you've created, and demonstrate it, by encoding and then decoding the implementation (the source code) of the Caesar cipher using Caesar cipher:
  * Encode the Python code
  * Decode the result of encoding
  * Compare the original and the recovered code  - you can add a box where you print the results — they should be the same.

Q1.2A: Implement function `ceasar_cipher_decode`, which will restore the cyphertext created by `ceasar_cipher` to plaintext
<pre>
------------
Input parameters:
- cyphertext - string, the encoded text that needs to be decyphered
- N - int, size of shift that was used in the encoding
------------
return value:
string, the text after decoding
</pre>

<!-- grade_id: cell-fd561d08a5195124, points: 0, grade: false, solution: true -->
```python
def ceasar_cipher_decode(ciphertext, shift_amount):
    """
    Decrypts the given ciphertext using the Caesar cipher algorithm with the specified shift amount.

    Args:
    ciphertext (str): The text to be decrypted.
    shift_amount (int): The number of positions each character in the ciphertext will be shifted back.

    Returns:
    str: The decrypted text (plaintext).
    """
    plaintext = ""
    alphabet = "abcdefghijklmnopqrstuvwxyz"

    for char in ciphertext:
        if char.lower() in alphabet:
            original_index = (alphabet.find(char.lower()) - shift_amount) % len(alphabet)
            plaintext += alphabet[original_index]
        else:
            plaintext += char

    return plaintext


```

<!-- grade_id: cell-a3e55a24e9359804, points: 8, grade: true, solution: false -->
```python
# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q1.2A --- Test your implementation:
# ---------------------------
print ("Testing the implementation of the 'ceasarCipherDecode' function..\n ")
assert type(ceasar_cipher_decode("ABC",3))==str,'output type is incorrect, should be string'
print ("good job!\nTests passed. There may be hidden tests too")
```

Output:
```
Testing the implementation of the 'ceasarCipherDecode' function..
 
good job!
Tests passed. There may be hidden tests too
```

Q1.2B: Implement function `return_encrypted_ceasar_code`, which will return a string that equals the code of `ceasar_cipher` after encryption, with N=2 . Then, implement a function `return_decoded_ceasar_code` which will call `return_encrypted_ceasar_code` and return the decrypted outcome.
<pre>
return_encrypted_ceasar_code()
------------
Input parameters:
none
------------
return value:
string, the encrypted source code
</pre>

<pre>
return_decoded_ceasar_code()
------------
Input parameters:
encrypted_code - the output of return_encrypted_ceasar_code()
------------
return value:
string, the source code restored to plaintext
</pre>

<!-- grade_id: cell-0c2fa8c5fcfb1b21, points: 0, grade: false, solution: true -->
```python
def return_encrypted_ceasar_code():
    """
    Returns the encrypted version of a predefined Caesar cipher encryption function's code.
    
    Returns:
    str: The encrypted code.
    """
    plaintext = (
        "def ceasar_cipher(plaintext, N):\n"
        "    ciphertext = \"\"\n"
        "    alphabet = \"abcdefghijklmnopqrstuvwxyz\"\n"
        "    for char in plaintext:\n"
        "        if char.lower() in alphabet:\n"
        "            idx = alphabet.find(char.lower())\n"
        "            new_idx = (idx + N) % len(alphabet)\n"
        "            ciphertext += alphabet[new_idx]\n"
        "        else:\n"
        "            ciphertext += char\n"
        "    return ciphertext"
    )
    ciphertext = ceasar_cipher(plaintext, 2)
    return ciphertext

def return_decoded_ceasar_code(encrypted_code=None):
    """
    Returns the decrypted version of the provided encrypted code using the Caesar cipher algorithm.
    
    Args:
    encrypted_code (str): The encrypted code to be decrypted. If not provided, the function encrypts a predefined code and then decrypts it.

    Returns:
    str: The decrypted code.
    """
    if encrypted_code is None:
        encrypted_code = return_encrypted_ceasar_code()
    decrypted_code = ceasar_cipher_decode(encrypted_code, 2)
    return decrypted_code

```

<!-- grade_id: cell-290384e75f4cb72f, points: 2, grade: true, solution: false -->
```python
# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q1.2B --- Test your implementation:
# ---------------------------
print ("Testing the implementation of the 'ceasarCipherDecode' function..\n ")
assert type(return_encrypted_ceasar_code())==str,'output type is incorrect, should be string'
assert type(return_decoded_ceasar_code())==str,'output type is incorrect, should be string'
assert type(return_decoded_ceasar_code("some_arbitrary_string"))==str,'output type is incorrect, should be string'
print ("good job!\nTests passed. There may be hidden tests too")
```

Output:
```
Testing the implementation of the 'ceasarCipherDecode' function..
 
good job!
Tests passed. There may be hidden tests too
```

3. Implement the encoder for the Vigenère cipher: https://en.wikipedia.org/wiki/Vigen%C3%A8re_cipher
(encode only the letters, leave other characters as-is).

Q1.3: Implement function `vigenere_encode`, which will encrypt the input text according to the vigenère code with the provided key.
<pre>
------------
Input parameters:
- plaintext - string, the text that needs to be encoded
- key - str
------------
return value:
string, the text after encoding
</pre>

<!-- grade_id: cell-a74f19da20d8c32c, points: 0, grade: false, solution: true -->
```python
def vigenere_encode(plaintext, key):
    """
    Encrypts the given plaintext using the Vigenère cipher algorithm with the specified key.
    
    Args:
    plaintext (str): The text to be encrypted.
    key (str): The encryption key.

    Returns:
    str: The encrypted text (ciphertext).
    """
    ciphertext = ""
    alphabet = "abcdefghijklmnopqrstuvwxyz"
    key_index = 0

    for char in plaintext:
        if char.lower() in alphabet:
            key_char = key[key_index % len(key)].lower()
            key_index += 1
            key_idx = alphabet.find(key_char)
            char_idx = alphabet.find(char.lower())
            new_idx = (key_idx + char_idx) % len(alphabet)
            if char.isupper():
                ciphertext += alphabet[new_idx].upper()
            else:
                ciphertext += alphabet[new_idx]
        else:
            ciphertext += char

    return ciphertext

```

<!-- grade_id: cell-fede5aa2fb49e35f, points: 10, grade: true, solution: false -->
```python
# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q1.3 --- Test your implementation:
# ---------------------------
print ("Testing the implementation of the 'vigenere_encode' function..\n ")
assert type(vigenere_encode("ABCabcABC","vig"))==str,'output type is incorrect, should be string'
print ("good job!\nTests passed. There may be hidden tests too")

```

Output:
```
Testing the implementation of the 'vigenere_encode' function..
 
good job!
Tests passed. There may be hidden tests too
```


4. Implement a decoder, and demonstrate by encoding and then decoding the implementation of the Vigenère cipher, as before.

<b>Please note:</b> in the case where the encoded text length isn't a multiple of the key, the decoder should still be able to reverse the encoding correctly! Make sure you've handled this case.

Q1.4A: Implement function `vigenere_decode`, which will restore the cyphertext created by `vigenere_encode` to plaintext
<pre>
------------
Input parameters:
- ciphertext - string, the encoded text that needs to be decrypted
- key - str
------------
return value:
string, the text after decoding
</pre>

<!-- grade_id: cell-1780c42a6fe02cf5, points: 0, grade: false, solution: true -->
```python
def vigenere_decode(encoded_text, decryption_key):
    """
    Decrypts the given ciphertext using the Vigenère cipher algorithm with the specified key.
    
    Args:
    encoded_text (str): The text to be decrypted.
    decryption_key (str): The decryption key.

    Returns:
    str: The decrypted text (plaintext).
    """
    decoded_text = ""
    alphabet = "abcdefghijklmnopqrstuvwxyz"
    key_position = 0

    for character in encoded_text:
        if character.lower() in alphabet:
            current_key_char = decryption_key[key_position % len(decryption_key)].lower()
            key_position += 1
            key_index = alphabet.find(current_key_char)
            char_index = alphabet.find(character.lower())
            new_index = (char_index - key_index) % len(alphabet)
            if character.isupper():
                decoded_text += alphabet[new_index].upper()
            else:
                decoded_text += alphabet[new_index]
        else:
            decoded_text += character

    return decoded_text

```

<!-- grade_id: cell-7525afb85c29fb34, points: 8, grade: true, solution: false -->
```python
# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q1.4A --- Test your implementation:
# ---------------------------
print ("Testing the implementation of the 'vigenere_decode' function..\n ")
assert type(vigenere_decode("ABCDefgh","vig"))==str,'output type is incorrect, should be string'
print ("good job!\nTests passed. There may be hidden tests too")
```

Output:
```
Testing the implementation of the 'vigenere_decode' function..
 
good job!
Tests passed. There may be hidden tests too
```

Q1.4B: Implement functions `return_encrypted_vigenere_code`, which will return a string that equals the code of `vigenere_encode` after encryption, with key="vig". Then, implement a function `return_decoded_vigenere_code` which will call `return_encrypted_vigenere_code` and return the decrypted outcome. You may print the output to verify that it matches the original.
<pre>
return_encrypted_vigenere_code()
------------
Input parameters:
none
------------
return value:
string, the encrypted source code
</pre>

<pre>
return_decoded_vigenere_code()
------------
Input parameters:
encrypted_code - the output of return_encrypted_vigenere_code()
------------
return value:
string, the source code restored to plaintext
</pre>

<!-- grade_id: cell-0e04f341a44584a7, points: 0, grade: false, solution: true -->
```python
def return_encrypted_vigenere_code():
    """
    Returns the encrypted version of a predefined Vigenère cipher encryption function's code.
    
    Returns:
    str: The encrypted code.
    """
    encode_func_str = (
        "def vigenere_encode(plaintext, key):\n"
        "    ciphertext = \"\"\n"
        "    alphabet = \"abcdefghijklmnopqrstuvwxyz\"\n"
        "    key_index = 0\n"
        "    for char in plaintext:\n"
        "        if char.lower() in alphabet:\n"
        "            key_char = key[key_index % len(key)].lower()\n"
        "            key_index += 1\n"
        "            key_idx = alphabet.find(key_char)\n"
        "            char_idx = alphabet.find(char.lower())\n"
        "            new_idx = (key_idx + char_idx) % len(alphabet)\n"
        "            if char.isupper():\n"
        "                ciphertext += alphabet[new_idx].upper()\n"
        "            else:\n"
        "                ciphertext += alphabet[new_idx]\n"
        "        else:\n"
        "            ciphertext += char\n"
        "    return ciphertext"
    )
    encrypted_code = vigenere_encode(encode_func_str, 'vig')
    return encrypted_code

def return_decoded_vigenere_code(encrypted_code=return_encrypted_vigenere_code()):
    """
    Returns the decrypted version of the provided encrypted code using the Vigenère cipher algorithm.
    
    Returns:
    str: The decrypted code.
    """
    decoded_code = vigenere_decode(encrypted_code, 'vig')
    return decoded_code

```

<!-- grade_id: cell-bef06a38f2c90ce7, points: 2, grade: true, solution: false -->
```python
# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q1.4B --- Test your implementation:
# ---------------------------
print ("Testing the implementation of the 'return_encrypted/return_decoded' function..\n ")
assert type(return_encrypted_vigenere_code())==str,'output type is incorrect, should be string'
assert type(return_decoded_vigenere_code())==str,'output type is incorrect, should be string'
assert type(return_decoded_vigenere_code("some_arbitrary_string"))==str,'output type is incorrect, should be string'
print ("good job!\nTests passed. There may be hidden tests too")
```

Output:
```
Testing the implementation of the 'return_encrypted/return_decoded' function..
 
good job!
Tests passed. There may be hidden tests too
```

# Question 2: Hash tables

The code cell below contains emergency phone numbers of the university (taken from https://in.bgu.ac.il/logistics/Pages/dialing_rules.aspx). A single entity may have several phone numbers.

You will implement a ‘truecaller’ facility: given a phone number, output the name of the entity. For example,
```python
caller = truecaller(t, 4100)
print(f"{caller} called!")
```
```truecaller``` should return “Police”.

if the number is not found, the function should return None

Q1. Implement 'truecaller' using a Python dictionary


```python
phonebook = [
    ["Security", '61888', '61555', '61444'],
    ["Police", '100', '4100'],
    ["Magen David Adom", '101', '4101'],
    ["Fire Brigade", '102', '4102'],
    ["Soroka Hospital", '4007']
]

def truecaller(phonebook, phone_number):
    for entry in phonebook:
        contact_name = entry[0]
        if str(phone_number) in entry[1:]:
            return contact_name
    return None

```

Q2.1A : Implement an auxiliary `make_trucaller_python_dictionary` function that converts the provided `phonebook` list into a Python dictionary
<pre>
------------
Input parameters:
- phn - a Python list, as `phonebook`, where every item is a list of strings. The first string is the caller's name, the following strings are phone numbers
------------
return value:
Python dictionary
------------
</pre>

<!-- grade_id: cell-87a450fdc22b0d47, points: 0, grade: false, solution: true -->
```python
def make_truecaller_python_dictionary(phonebook_list=phonebook):
    """
    Converts a phonebook list into a dictionary for efficient lookup.

    Args:
    phonebook_list (list): A list of contacts where each contact is a list of a name followed by phone numbers.

    Returns:
    dict: A dictionary where keys are phone numbers and values are contact names.
    """
    phonebook_dict = {}
    for contact in phonebook_list:
        contact_name = contact[0]
        phone_numbers = contact[1:]
        for number in phone_numbers:
            phonebook_dict[number] = contact_name
    return phonebook_dict


```

Q2.1B : Implement a `truecaller1` function, relying on `make_truecaller_python_dictionary`, which takes as input a phone number, and outputs a string with the caller's name as it appears in the phonebook (or returns None if the number couldn't be found).
<pre>
------------
Input parameters:
- phone_number - a string, containing a phone number
------------
return value:
a string , with the caller's name
------------
</pre>

<!-- grade_id: cell-9e86c1484dcf36f4, points: 0, grade: false, solution: true -->
```python
phonebook_dict = make_truecaller_python_dictionary(phonebook)

def truecaller1(phone_number, caller_dict=phonebook_dict):
    if phone_number in caller_dict:
        return caller_dict[phone_number]
    else:
        return None

```

<!-- grade_id: cell-1be6975be3ec4742, points: 12, grade: true, solution: false -->
```python
# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q2.1 --- Test your implementation:
# ---------------------------
print ("Testing the implementation of the 'make_truecaller_python_dictionary' function..\n ")
assert type(make_truecaller_python_dictionary())==dict,'output type is incorrect, should be python dictionary'
print ("Testing the implementation of the 'truecaller1' function..\n ")
assert type(truecaller1('101'))==str,'output type is incorrect, should be string'
print ("good job!\nTests passed. There may be hidden tests too")

```

Output:
```
Testing the implementation of the 'make_truecaller_python_dictionary' function..
 
Testing the implementation of the 'truecaller1' function..
 
good job!
Tests passed. There may be hidden tests too
```


Q2. Implement "truecaller2" using your own implementation of dictionary:

<!-- grade_id: cell-f06f7223c2e94a50, points: 0, grade: false, solution: true -->
```python
#You can add auxiliary functions, according to your discretion
# YOUR CODE HERE

```

Q2.2A : Implement an auxiliary `make_truecaller_own_implement_dictionary` function that converts the provided `phonebook` list into a structure that behaves like a hash table
    
<pre>
------------
Input parameters:
- phn - a Python list, as `phonebook`, where every item is a list of strings. The first string is the caller's name, the following strings are phone numbers
------------
return value:
a Python list
------------
</pre>

<!-- grade_id: cell-0f5cd96a463d46c9, points: 0, grade: false, solution: true -->
```python
def make_truecaller_own_implement_dictionary(phonebook_list=phonebook):
    hash_table = [None] * 101  # Initialize hash table with 101 buckets (a prime number)

    for entry in phonebook_list:
        contact_name = entry[0]
        for number in entry[1:]:
            hash_value = sum([ord(char) for char in number]) % 101  # Calculate hash value for phone number
            if hash_table[hash_value] is None:
                hash_table[hash_value] = [[number, contact_name]]  # Create new list if bucket is empty
            else:
                hash_table[hash_value].append([number, contact_name])  # Add entry to existing list in bucket

    return hash_table

```

Q2.2B : Implement a `truecaller2` function, relying on `make_truecaller_own_implement_dictionary`, which takes as input a phone number, and outputs a string with the caller's name as it appears in the phonebook (or returns None if the number couldn't be found).
<pre>
------------
Input parameters:
- phone_number - a string, containing a phone number
------------
return value:
a string , with the caller's name
------------
</pre>

<!-- grade_id: cell-a2c907715447ebae, points: 0, grade: false, solution: true -->
```python
custom_phonebook_dict = make_truecaller_own_implement_dictionary()

def truecaller2(phone_number, phonebook_dict=custom_phonebook_dict):
    """
    Looks up the contact name for a given phone number using a custom hash table implementation.

    Args:
    phone_number (str): The phone number to look up.
    phonebook_dict (list): The custom hash table dictionary.

    Returns:
    str: The contact name if the phone number is found, otherwise None.
    """
    hash_value = sum(ord(char) for char in phone_number) % 101  # Calculate hash value for phone number
    if phonebook_dict[hash_value] is None:
        return None
    else:
        for entry in phonebook_dict[hash_value]:
            if entry[0] == phone_number:
                return entry[1]
        return None


```

<!-- grade_id: cell-9bd7a35c7e4305a6, points: 12, grade: true, solution: false -->
```python
# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q2.2 --- Test your implementation:
# ---------------------------
print ("Testing the implementation of the 'make_truecaller_python_dictionary' function..\n ")
assert type(make_truecaller_own_implement_dictionary())==list,'output type is incorrect, should be list'
print ("Testing the implementation of the 'truecaller1' function..\n ")
assert type(truecaller2('102'))==str,'output type is incorrect, should be string'
print ("good job!\nTests passed. There may be hidden tests too")

```

Output:
```
Testing the implementation of the 'make_truecaller_python_dictionary' function..
 
Testing the implementation of the 'truecaller1' function..
 
good job!
Tests passed. There may be hidden tests too
```


3. Implement a hash function `hash_string` for ids. Treat ids as _strings_ of characters.


<pre>
------------
Input parameters:
- str_to_hash - a string
------------
return value:
an integer
------------
</pre>

<!-- grade_id: cell-907e0df43371898d, points: 0, grade: false, solution: true -->
```python
def hash_string(string_to_hash):
    """
    Computes the hash value of a given string using a basic hash function.

    Args:
    string_to_hash (str): The string to hash.

    Returns:
    int: The computed hash value.
    """
    hash_value = 0
    for char in string_to_hash:
        hash_value = (hash_value * 31 + ord(char)) % 10000  # Compute hash value using a multiplier and modulo operation
    return hash_value
```

<!-- grade_id: cell-6074fd7cee4c2b6d, points: 12, grade: true, solution: false -->
```python

# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q2.3 --- Test your implementation:
# ---------------------------
print ("Testing the implementation of the 'hash_string' function..\n ")
assert type(hash_string("some string"))==int,'output should be an integer'
print ("good job!\nTests passed. There may be hidden tests too")


```

Output:
```
Testing the implementation of the 'hash_string' function..
 
good job!
Tests passed. There may be hidden tests too
```


4. A check digit is used in ids in order to detect errors and make faking an id a bit harder. The exact formula for ids' check digit can be found in https://he.wikipedia.org/wiki/%D7%A1%D7%A4%D7%A8%D7%AA_%D7%91%D7%99%D7%A7%D7%95%D7%A8%D7%AA. Implement a function `check_id` which checks if a given id is legal with 2 conditions:
1. It is of length 9
2. The check digit is correct

Implement the function `check_id` that returns whether an id is legal or not:

<pre>
------------
Input parameters:
- id: a string
------------
return value:
- boolean
------------
</pre>

<!-- grade_id: cell-a5389da6f18d38bd, points: 0, grade: false, solution: true -->
```python
#You can add auxiliary functions, according to your discretion
# YOUR CODE HERE

```

<!-- grade_id: cell-cba9f3c0fcdfb195, points: 0, grade: false, solution: true -->
```python
def check_id(identifier):
    """
    Validates an ID number using a weighted checksum algorithm.

    Args:
    identifier (str): The ID number to validate.

    Returns:
    bool: True if the ID is valid, False otherwise.
    """
    if len(identifier) != 9:
        return False

    weights = [1, 2, 1, 2, 1, 2, 1, 2, 1]
    total_sum = 0

    for i in range(9):
        digit = int(identifier[i])
        product = weights[i] * digit
        if product > 9:
            product = product // 10 + product % 10  # Sum the digits of the product
        total_sum += product

    return total_sum % 10 == 0

```

<!-- grade_id: cell-cf67edd4a8ae7aa5, points: 14, grade: true, solution: false -->
```python


# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q2.4 --- Test your implementation:
# ---------------------------
print ("Testing the implementation of the min_size- functions..\n ")
assert check_id("123456789")==False,'output should be False'
assert check_id("123456782")==True,'output should be True'
print ("good job!\nTests passed. There may be hidden tests too")

```

Output:
```
Testing the implementation of the min_size- functions..
 
good job!
Tests passed. There may be hidden tests too
```

<!-- grade_id: cell-cf67edd4a8ae7aa5, points: 14, grade: true, solution: false -->
5. Hashing is called ‘perfect’ if there are no ‘key conflicts‘. Can you find if there is a value that has a conflict with your own id, when using the hash function you implemented in Q2.3?

<!-- grade_id: cell-cf67edd4a8ae7aa5, points: 14, grade: true, solution: false -->
Implement the function `find_conflict` that returns if there is an id with the same hash value as your own id, or an empty string otherwise:

<pre>
------------
Input parameters:
- id: your own id
------------
return value:
- string, either empty or a legal id with the same hash value as your own id
------------
</pre>

<!-- grade_id: cell-cba9f3c0fcdfb195, points: 0, grade: false, solution: true -->
```python

def find_conflict(identifier):
    """
    Finds a conflicting ID that has the same hash value as the given ID.

    Args:
    identifier (str): The ID to find a conflict for.

    Returns:
    str: A conflicting ID with the same hash value or an empty string if no conflict is found.
    """
    hash_value = hash_string(identifier)
    # Generate a list of all legal IDs and check if any of them have the same hash value
    for i in range(1000000000):
        candidate_id = str(i).zfill(9)  # Pad candidate ID with leading zeros
        if check_id(candidate_id) and hash_string(candidate_id) == hash_value and candidate_id != identifier:
            return candidate_id
    return ""

```